# BT03: Pandas

Họ tên: ...

MSSV: ...

---

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells` để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - Thư mục `BT03`
        - File `BT03-Pandas.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. Đuôi của file nén phải là `.zip` (chứ không được `.rar` hay cái khác). <font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài này.</font>

---

## Import

```python
%matplotlib inline
import matplotlib.pyplot as plt
import pandas as pd
import numpy as np
# YOUR CODE HERE (OPTION) 
# Nếu cần các thư viện khác thì bạn có thể import ở đây
```

---

## Lấy dữ liệu

Dữ liệu được dùng trong bài tập này là file "mvt.csv" đính kèm; đây là dữ liệu về tội phạm ăn cắp xe ở thành phố Chicago, bang Illinois, Hoa Kỳ ("mvt" là viết tắt của "motor vehicle theft"). Dữ liệu này được lấy từ khóa học "The Analytics Edge" ở trang edx.org, và gốc của dữ liệu là [ở đây](https://data.cityofchicago.org/Public-Safety/Crimes-2001-to-present/ijzp-q8t2).

---

## Khám phá dữ liệu

### Đọc dữ liệu từ file csv vào dataframe

```python
mvt_df = pd.read_csv("mvt.csv")
mvt_df.head()
```

### Dữ liệu gồm có bao nhiêu dòng và bao nhiêu cột? (0.25đ)

Bạn tính số dòng và số cột (không được code cứng con số) và lần lượt lưu vào 2 biến là `num_rows` và `num_cols`. 

<!-- grade_id: cell-bc8810dbc4930a46, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41d8eb768a571d35, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert num_rows == 191641
assert num_cols == 11
```

### Mỗi dòng và mỗi cột có ý nghĩa gì?

Dưới đây là phần mô tả của khóa học "The Analytics Edge" về các cột trong file "mvt.csv":
- **ID**: a unique identifier for each observation.
- **Date**: the date the crime occurred.
- **LocationDescription**: the location where the crime occurred.
- **Arrest**: whether or not an arrest was made for the crime (TRUE if an arrest was made, and FALSE if an arrest was not made).
- **Domestic**: whether or not the crime was a domestic crime, meaning that it was committed against a family member (TRUE if it was domestic, and FALSE if it was not domestic).
- **Beat**: the area, or "beat" in which the crime occurred. This is the smallest regional division defined by the Chicago police department.
- **District**: the police district in which the crime occured. Each district is composed of many beats, and are defined by the Chicago Police Department.
- **CommunityArea**: the community area in which the crime occurred. Since the 1920s, Chicago has been divided into what are called "community areas", of which there are now 77. The community areas were devised in an attempt to create socially homogeneous regions.
- **Year**: the year in which the crime occurred.
- **Latitude**: the latitude of the location at which the crime occurred.
- **Longitude**: the longitude of the location at which the crime occurred.

Nhìn vào dữ liệu ở `mvt_df`, ta có thể thấy là mỗi dòng ứng với thông tin của một vụ phạm tội.

### Mỗi cột  có dtype là gì? (0.25đ)

Bạn lấy `dtype` (kiểu dữ liệu của mỗi phần tử) của mỗi cột trong `mvt_df` và lưu kết quả vào series `col_dtypes`; series này có index là tên cột. 

<!-- grade_id: cell-e7c33d274bed2b0c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84a2c0629f12e655, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert col_dtypes['ID'] == int
assert col_dtypes['Beat'] == int
assert col_dtypes['Year'] == int
assert col_dtypes['Latitude'] == float
assert col_dtypes['Longitude'] == float
assert col_dtypes['Date'] == object
assert col_dtypes['LocationDescription'] == object
assert col_dtypes['District'] == object
assert col_dtypes['CommunityArea'] == object
assert col_dtypes['Arrest'] == bool
assert col_dtypes['Domestic'] == bool
```

### Cột có dtype là object nghĩa là sao? (0.5đ)

`mvt_df` có 4 cột có `dtype` là `object`: "Date", "LocationDescription", "District", "CommunityArea". Trong Pandas, kiểu dữ liệu `object` thường ám chỉ chuỗi, nhưng thật ra kiểu dữ liệu `object` có thể chứa một đối tượng bất kỳ trong Python (vì thật ra ở bên dưới kiểu dữ liệu `object` chứa địa chỉ). Nếu một cột trong dataframe có `dtype` là `object` thì có thể các phần tử trong cột này sẽ có kiểu dữ liệu khác nhau; để biết được kiểu dữ liệu thật sự của các phần tử trong cột này thì ta phải truy xuất vào từng phần tử và dùng câu lệnh `type` để xem kiểu dữ liệu, ví dụ `type(s.iloc[i])` với `s` là một cột. Lúc đọc dữ liệu từ file "mvt.csv" lên, nếu bạn để ý thì sẽ thấy warning là: "Columns (6,7) have mixed types", nghĩa là cột 6 ("District") và 7 ("CommunityArea") có `dtype` không đồng nhất (trong nội bộ mỗi cột). Ta muốn xem thử trong nội bộ mỗi cột này có các kiểu dữ liệu nào.

Công việc của bạn là viết hàm nhận đầu vào là một cột `s` (series) có `dtype` là `object` và trả về tập các kiểu dữ liệu thật sự có trong cột này.

<!-- grade_id: cell-9c4658645a61bd6e, points: 0, grade: false, solution: true -->
```python
def open_object_dtype(s):
    dtypes = set()
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return dtypes
```

<!-- grade_id: cell-f767e07e65480668, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert open_object_dtype(mvt_df['Date']) == {str}
assert open_object_dtype(mvt_df['LocationDescription']) == {str}
assert open_object_dtype(mvt_df['District']) == {str, float}
assert open_object_dtype(mvt_df['CommunityArea']) == {str, float}
```

**<font color=red>Vấn đề:</font>**

- Cột "Date" đang có dtype là `str`. Để có thể tiếp tục khám phá thêm về cột này, ta sẽ thực hiện bước tiền xử lý là chuyển sang dạng `datetime`.
- Cột "District" và "CommunityArea" có dtype gồm 2 loại là `str` và `float`. Nếu thí nghiệm thêm thì bạn sẽ thấy `float` là ứng với `NaN`, còn `str` là ứng với mã (dạng số) của "District" hoặc "CommunityArea". Tuy nhiên, cũng có trường hợp mã là `float` (ví dụ, cột "CommunityArea", index 131072). Để có thể tiếp tục khám phá thêm về 2 cột này, ta cũng sẽ thực hiện bước tiền xử lý là chuyển tất cả sang dtype là `float`.

---

## Tiền xử lý

### Chuyển dtype của cột "Date" sang datetime (0.5đ)

Bạn hãy chuyển `dtype` của cột "Date" sang kiểu dữ liệu thời gian (sau khi làm xong thì cột "Date" trong `mvt_df` phải thay đổi `dtype`). Gợi ý: bạn có thể dùng câu lệnh `pd.to_datetime`; khi dùng `pd.to_datetime`, bạn truyền tham số `format="%m/%d/%y %H:%M"` (nếu không thì sẽ chạy rất chậm vì Pandas phải đoán format).

<!-- grade_id: cell-39a0a94d173359e5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8d966ec31928fc64, points: 0.5, grade: true, solution: false -->
```python
# TEST 
# Nếu dtype của cột "Date" đã được chuyển sang datetime thì phải dùng được .dt.
mvt_df['Date'].dt.year.min() == 2001 
```

### Chuyển dtype của cột "District" và "CommunityArea" sang float (0.75đ)

Bạn hãy chuyển `dtype` của cột "District" và "CommunityArea" sang `float` (sau khi làm xong thì cột "District" và "CommunityArea" trong `mvt_df` phải thay đổi `dtype`). Nếu bạn thử chuyển thì sẽ bị báo lỗi vì trong các cột này có giá trị là " " (chuỗi có một khoảng trắng). Vậy thì trước khi chuyển, bạn phải replace các giá trị " " bằng NaN.  

<!-- grade_id: cell-82a6fc90114be977, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-02292b1dd9159219, points: 0.75, grade: true, solution: false -->
```python
# TEST
assert mvt_df['District'].dtype == float
assert mvt_df['CommunityArea'].dtype == float
```

---

## Quay lại bước khám phá dữ liệu

### Các cột dạng numeric và datetime có miền giá trị như thế nào? (0.5đ)

```python
mvt_df.info()
```

Hiện có 8 cột có dạng numeric hoặc datetime: "ID", "Date", "Beat", "District", "CommunityArea", "Year", "Latitude", "Longitude". Tuy nhiên, cột "ID" thì không có gì đáng quan tâm; 3 cột "Beat", "District", "CommunityArea" thật ra là dạng categorical (vì là mã được đánh dưới dạng số). Do đó, ở đây, bạn chỉ tính miền giá trị (min, max) của 4 cột là "Date", "Year", "Latitude", "Longitude". Bạn lưu kết quả vào dataframe `col_ranges_df`. Dataframe này có 2 dòng là "min", "max"; và có 4 cột là "Date", "Year", "Latitude", "Longitude".  

<!-- grade_id: cell-4c64c2ce1eea9f1d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

col_ranges_df
```

<!-- grade_id: cell-d148b332b71aabcc, points: 0.5, grade: true, solution: false -->
```python
## TEST
assert col_ranges_df.loc['min', 'Date'] == pd.Timestamp('2001-01-01 00:01:00')
assert col_ranges_df.loc['max', 'Date'] == pd.Timestamp('2012-12-31 23:15:00')
assert col_ranges_df.loc['min', 'Year'] == 2001
assert col_ranges_df.loc['max', 'Year'] == 2012
assert f"{col_ranges_df.loc['min', 'Latitude']:.3f}" == '41.645'
assert f"{col_ranges_df.loc['max', 'Latitude']:.3f}" == '42.023'
assert f"{col_ranges_df.loc['min', 'Longitude']:.3f}" == '-87.926'
assert f"{col_ranges_df.loc['max', 'Longitude']:.3f}" == '-87.525'
```

### Các cột dạng categorical có các giá trị nào? (1đ)

```python
mvt_df.info()
```

`mvt_df` có 4 cột dạng categorical là "LocationDescription", "Beat", "District", "CommunityArea". Bạn hãy tính "range" của mỗi cột này: số lượng các giá trị khác nhau (không tính NaN), và list gồm 3 giá trị khác nhau đầu tiên sau khi đã sort tăng dần tất cả các giá trị khác nhau. Bạn lưu kết quả vào dataframe `cat_col_ranges_df`. Dataframe này có 2 dòng là "NumUniqueVals", "ThreeFirstVals"; và có 4 cột là "LocationDescription", "Beat", "District", "CommunityArea".

<!-- grade_id: cell-22007ecdcdfa0acc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

cat_col_ranges_df
```

<!-- grade_id: cell-a4cbf35ddac608bb, points: 1, grade: true, solution: false -->
```python
# TEST
df = cat_col_ranges_df
assert df.loc['NumUniqueVals', 'LocationDescription'] == 78
assert df.loc['ThreeFirstVals', 'LocationDescription'] == ['ABANDONED BUILDING', 
                                                           'AIRPORT BUILDING NON-TERMINAL - NON-SECURE AREA', 
                                                           'AIRPORT BUILDING NON-TERMINAL - SECURE AREA']
assert df.loc['NumUniqueVals', 'Beat'] == 300
assert df.loc['ThreeFirstVals', 'Beat'] == [111, 112, 113]
assert df.loc['NumUniqueVals', 'District'] == 26
assert df.loc['ThreeFirstVals', 'District'] == [1.0, 2.0, 3.0]
assert df.loc['NumUniqueVals', 'CommunityArea'] == 78
assert df.loc['ThreeFirstVals', 'CommunityArea'] == [0.0, 1.0, 2.0]
```

### Làm rõ hơn về ý nghĩa của các cột "Beat", "District", "CommunityArea" (1đ)

Theo mô tả của khóa học "The Analytics Edge":
- **Beat**: the area, or "beat" in which the crime occurred. This is the smallest regional division defined by the Chicago police department.
- **District**: the police district in which the crime occured. Each district is composed of many beats, and are defined by the Chicago Police Department.
- **CommunityArea**: the community area in which the crime occurred. Since the 1920s, Chicago has been divided into what are called "community areas", of which there are now 77. The community areas were devised in an attempt to create socially homogeneous regions.

Như vậy, "Beat" là đơn vị nhỏ nhất, "District" thì gồm nhiều "Beat". Vậy có phải là các "District" không bị overlap với nhau không? (không có trường hợp mà một "Beat" vừa thuộc "District" này, vừa thuộc "District" kia). Và mỗi "CommunityArea" là gồm nhiều "Beat" và cũng không overlap với nhau? Để trả lời cho hai câu hỏi này, dưới đây bạn sẽ tính tỉ lệ (từ 0 đến 1) các "Beat" mà thuộc nhiều hơn một "District", và tỉ lệ các "Beat" mà thuộc nhiều hơn một "CommunityArea". Bạn lần lượt lưu các kết quả vào hai biến là `multi_district_ratio` và `multi_comarea_ratio`.

<!-- grade_id: cell-f312e85b9123e0b0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6faf22732c2063c1, points: 1, grade: true, solution: false -->
```python
# TEST
assert f'{multi_district_ratio:.3f}' == '0.210'
assert f'{multi_comarea_ratio:.3f}' == '0.603'
```

### Mỗi cột có bao nhiêu giá trị thiếu? (0.5đ)

Bạn hãy tính số lượng giá trị thiếu của mỗi cột và lưu kết quả vào series `col_num_nans`, trong đó index là tên cột. 

<!-- grade_id: cell-b455eecae88a7e3a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-065d5a19134a5c73, points: 0.5, grade: true, solution: false -->
```python
# TEST
s = col_num_nans
assert len(s[s > 0]) == 4
assert s.loc['District'] == 43056
assert s.loc['CommunityArea'] == 24616
assert s.loc['Latitude'] == 2276
assert s.loc['Longitude'] == 2276
```

---

## Đưa ra các câu hỏi cần trả lời

Sau khi đã khám phá dữ liệu và hiểu hơn về dữ liệu, ta thấy có một số câu hỏi có thể được trả lời bằng dữ liệu:
- Số lượng các vụ phạm tội thay đổi như thế nào theo năm?
- Với các vụ án xảy ra ở mỗi năm, tỉ lệ bắt được thủ phạm tính đến thời điểm dữ liệu được thu thập là bao nhiêu?
- Trong mỗi năm, "Beat" nào có nhiều vụ phạm tội nhất và "Beat" nào có ít vụ phạm tội nhất?
- Trong 3 năm gần đây, 5 loại địa điểm nào (ngoài đường, chỗ đậu xe, ...) có nhiều vụ phạm tội nhất?
- Trong mỗi năm của 3 năm gần đây, số lượng các vụ phạm tội thay đổi như thế nào theo tháng?

Trong đồ án cuối kỳ, bạn sẽ phải suy nghĩ để đưa ra các câu hỏi có ý nghĩa và cũng phải giải thích để người đọc thấy được là câu hỏi có ý nghĩa. Trong bài tập này, mục đích chính là tập luyện các thao tác xử lý dữ liệu của Pandas nên mình không cần giải thích để người đọc thấy được là các câu hỏi có ý nghĩa. May quá 😉. 

---

## Tiền xử lý

### Thêm cột "Month" vào dataframe (0.5đ)

Trong các câu hỏi ở trên, có câu cần đến thông tin về tháng của vụ phạm tội. Do đó, trước khi phân tích để trả lời cho các câu hỏi, bạn hãy tiền xử lý bằng cách thêm cột "Month" (tháng xảy ra vụ phạm tội) vào `mvt_df`. Sau khi thêm thì `mvt_df` phải thay đổi.

<!-- grade_id: cell-b11a3d54448e8905, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6a301e064922ea1e, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert mvt_df['Month'].min() == 1
assert mvt_df['Month'].max() == 12
```

---

## Phân tích dữ liệu để trả lời cho các câu hỏi 

### Số lượng các vụ phạm tội thay đổi như thế nào theo năm? (0.75đ)

Để trả lời cho câu hỏi này, bạn sẽ làm như sau:
- Bước 1: Tính số vụ phạm tội của mỗi năm. Bạn lưu kết quả vào series `num_thefts_each_year`, trong đó index là năm.
- Bước 2: Từ kết quả ở trên, bạn vẽ đồ thị dạng line, trong đó trục hoành là năm và trục tung là số lượng vụ phạm tội. Bạn dùng thêm marker là điểm tròn để cho thấy rõ các điểm dữ liệu. Bạn đặt tên trục hoành là "Year" và tên trục tung là "NumThefts".

Code bước 1. 

<!-- grade_id: cell-ca92bcde951cbcd9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-10c1e16b0a2b2bf8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert num_thefts_each_year.loc[2001] == 20669
assert num_thefts_each_year.loc[2012] == 14092
```

Code bước 2. 

<!-- grade_id: cell-c3f22e75a64dab00, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Hmm, ở Mỹ, năm 2008-2009 xảy ra cuộc khủng hoảng kinh tế, không biết có liên quan gì không ...

### Với các vụ án xảy ra ở mỗi năm, tỉ lệ bắt được thủ phạm tính đến thời điểm dữ liệu được thu thập là bao nhiêu? (1đ)

Để trả lời cho câu hỏi này, bạn sẽ làm như sau:
- Bước 1: Tính tỉ lệ (từ 0 đến 1) bắt được thủ phạm của mỗi năm (năm ở đây là năm xảy ra các vụ phạm tội). Bạn lưu kết quả vào series `arrest_ratio_each_year`, trong đó index là năm.
- Bước 2: Từ kết quả ở trên, bạn vẽ đồ thị dạng line, trong đó trục hoành là năm và trục tung là tỉ lệ bị bắt. Bạn dùng thêm marker là điểm tròn để cho thấy rõ các điểm dữ liệu. Bạn đặt tên trục hoành là "Year" và tên trục tung là "ArrestRatio". Ngoài ra, bạn cũng cần chỉnh miền giá trị của trục tung là từ 0 đến 1.

Code bước 1. 

<!-- grade_id: cell-129a07978e3c0de1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-81a06dfa2c96900c, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert f"{arrest_ratio_each_year.loc[2001]:.3f}" == '0.104'
assert f"{arrest_ratio_each_year.loc[2012]:.3f}" == '0.039'
```

Code bước 2. 

<!-- grade_id: cell-fe3eae5b8130e5ec, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Bạn sẽ thấy tỉ lệ bị bắt nhìn chung là giảm trong các năm gần đây. Tại sao lại như vậy? Nói chung là ta không biết chắc được. Nhưng bạn thử suy nghĩ và đưa ra một giả thuyết để giải thích cho xu hướng này dựa vào thông tin sau: năm ở đây là năm xảy ra vụ phạm tội và dữ liệu này được thu thập vào khoảng năm 2013 (năm mở của khóa học "The Analytics Edge").

<!-- grade_id: cell-d91e2d8fd4045ebb, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE

###  Trong mỗi năm, "Beat" nào có nhiều vụ phạm tội nhất và "Beat" nào có ít vụ phạm tội nhất? (0.75đ)

Bạn lưu kết quả vào dataframe `two_beats_each_year_df`. Dataframe này có 12 dòng ứng với 12 năm là 2001, 2002, ..., 2012; và có 2 cột là "SafeBeat" ("Beat" có ít vụ phạm tội nhất) và "DangerousBeat" ("Beat" có nhiều vụ phạm tội nhất).

<!-- grade_id: cell-7ef114a2b6ffaad2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

two_beats_each_year_df
```

<!-- grade_id: cell-7bab4fca5fec8531, points: 0.75, grade: true, solution: false -->
```python
# TEST
df = two_beats_each_year_df
assert df.loc[2001, 'SafeBeat'] == 134
assert df.loc[2001, 'DangerousBeat'] == 833
assert df.loc[2012, 'SafeBeat'] == 134
assert df.loc[2012, 'DangerousBeat'] == 824
```

### Trong 3 năm gần đây, 5 loại địa điểm nào (ngoài đường, chỗ đậu xe, ...) có nhiều vụ phạm tội nhất? (0.75đ)

Loại địa điểm ở đây được lưu trong cột "LocationDescription". Năm loại địa điểm trong câu hỏi là không tính "OTHER". Bạn lưu kết quả vào list `top5_locs` theo thứ tự phần tử đầu tiên là có nhiều vụ phạm tội nhất, phần tử thứ hai là có nhiều vụ phạm tội nhì, ...

<!-- grade_id: cell-9f05d5d157d6c417, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e91bd1e3d1aceec, points: 0.75, grade: true, solution: false -->
```python
# TEST
assert top5_locs == ['STREET', 
                     'PARKING LOT/GARAGE(NON.RESID.)', 
                     'RESIDENTIAL YARD (FRONT/BACK)', 
                     'VEHICLE NON-COMMERCIAL', 
                     'ALLEY']
```

### Trong mỗi năm của 3 năm gần đây, số lượng các vụ phạm tội thay đổi như thế nào theo tháng? (1đ)

Để trả lời cho câu hỏi này, bạn sẽ làm như sau:
- Bước 1: Tính dataframe `num_thefts_each_month_df` gồm có 3 cột (tên là 2010, 2011, 2012) ứng với số lượng vụ phạm tội theo các tháng của 3 năm, và 12 dòng (tên là 1, 2, ..., 12) ứng với 12 tháng.  
- Bước 2: Từ kết quả ở trên, bạn vẽ 3 đường dạng line (ứng với 3 năm) trên cùng một đồ thị, trong đó trục hoành là tháng và trục tung là số lượng vụ phạm tội. Bạn dùng thêm marker là điểm tròn để cho thấy rõ các điểm dữ liệu. Bạn đặt tên trục hoành là "Month" và tên trục tung là "NumThefts". Bạn cũng cần dùng legend để ghi chú là đường nào thì ứng với năm nào.

Code bước 1. 

<!-- grade_id: cell-8199894515a93739, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

num_thefts_each_month_df
```

<!-- grade_id: cell-1b37c024f3bf33e9, points: 0.75, grade: true, solution: false -->
```python
# TEST
df = num_thefts_each_month_df
assert df.shape == (12, 3)
assert df.loc[1, 2010] == 1278
assert df.loc[1, 2011] == 1508
assert df.loc[1, 2012] == 1253
```

Code bước 2. 

<!-- grade_id: cell-443fcd8b92392f58, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Hmm ...